# Lab 9 - Relational Algebra 

*This examples build on examples from Data100 and also Ch. 23 of the Data Science by Scratch book.* 

In this lab notebook you will **write functions to implement the basic relational algebra operators.**

## Lab Setup

```python
import numpy as np
import pandas as pd

import otter
grader = otter.Notebook()
```

# Huskies Sailing Club 

The Huskies Sailing Club wants to better organize their user data. Your first job is to implement code relational algebra operators in Python. 

You may want to refer to the [pandas](http://pandas.pydata.org/pandas-docs/stable/) documentation. 

Here are the Sailing Club's databases.  Your job implement a variety of unary and binary relational algebra operators. 

```python
young_sailors = pd.DataFrame({
        "sid":      [2701, 18869, 63940, 21869, 17436],
        "sname":    ["Grace", "Morgan", "Anne", "Jack", "Javier"],
        "rating":   [8, 6, 4, 9, 3],
        "age":      [25, 26, 21, 27, 22],
    })
salty_sailors = pd.DataFrame({
        "sid":      [2701, 17436, 45433, 22689, 46535],
        "sname":    ["Grace", "Javier", "Balon", "Euron", "Victarion"],
        "rating":   [8, 3, 7, 10, 2],
        "age":      [25, 22, 39, 35, 37],
    })
boats = pd.DataFrame({
        "bid":      [41116, 54505, 50041, 35168, 58324],
        "bname":    ["The Black Sparrow", "The Great Kraken", "The Prophetess", 
                     "Iron Victory", "Silence"],
        "color":    ["Black", "Orange", "Silver", "Red", "Grey"],
    })
reservations = pd.DataFrame({
        "sid":      [21869, 45433, 18869, 22689, 21869, 17436, 63940, 45433, 
                     21869, 18869],
        "bid":      [41116, 35168, 50041, 41116, 58324, 50041, 54505, 41116, 
                     50041, 41116],
        "day":      ["3/1", "3/1", "3/2", "3/2", "3/2", "3/3", "3/3", "3/3", 
                     "3/3", "3/4"],
    })
```

```python
young_sailors
```

Output:
```
sid   sname  rating  age
0   2701   Grace       8   25
1  18869  Morgan       6   26
2  63940    Anne       4   21
3  21869    Jack       9   27
4  17436  Javier       3   22
```

```python
salty_sailors
```

Output:
```
sid      sname  rating  age
0   2701      Grace       8   25
1  17436     Javier       3   22
2  45433      Balon       7   39
3  22689      Euron      10   35
4  46535  Victarion       2   37
```

```python
boats
```

Output:
```
bid              bname   color
0  41116  The Black Sparrow   Black
1  54505   The Great Kraken  Orange
2  50041     The Prophetess  Silver
3  35168       Iron Victory     Red
4  58324            Silence    Grey
```

```python
reservations
```

Output:
```
sid    bid  day
0  21869  41116  3/1
1  45433  35168  3/1
2  18869  50041  3/2
3  22689  41116  3/2
4  21869  58324  3/2
5  17436  50041  3/3
6  63940  54505  3/3
7  45433  41116  3/3
8  21869  50041  3/3
9  18869  41116  3/4
```

## Exercise 1: Projection

The first relational algebra function to implement is **Projection**.

Create a function `project` that takes as arguments a DataFrame and a list of columns to select and returns the projection (another DataFrame)

```python
def project(df, columns):     
    return df[columns]

df1a = project(salty_sailors, ["sname", "age"])
df1b = project(boats, ["color"])
```

Example test:  
The output for `project(salty_sailors, ["sname", "age"])` should be: 

|   | sname  |  age  | 
|---|--------|-------|
| 0 | Grace  | 25    |
| 1 | Javier | 22    | 
| 2 | Balon  | 39    | 
| 3 | Euron  | 35    | 
| 4 | Victarion | 37 |


```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Exercise 2: Selection 

Next, you will implement the **Selection** operator.

Create a function `select` that has arguments of a DataFrame and a condition (function) which determines which rows to select.  The function returns a DataFrame after performing the selection.  

For example: 
`good_sailors = select(young_sailors, lambda x: x["rating"] > 6)`

```python
def select(df, condition):     
    return df.loc[condition]

dfq2 = select(young_sailors, lambda x: x["rating"] > 6)
```

The output for `select(young_sailors, lambda x: x["rating"] > 6)` should be:

|   | sid   | sname  | rating | age |
|---|-------|--------|--------|-----|
| 0 | 2701  | Grace  | 8      | 25  |
| 3 | 21869 | Jack   | 9      | 27  |

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Exercise 3: Union 

Create a function, `union`, (a binary operator) that takes two DataFrames as input.  

You can assume the two dataframes are union compatible (that is, they have the same number of columns, and their columns have the same types).  

Return the result of the union operation.  

Assume set semantics, that is, duplicates should be removed. 

```python
def union(df1, df2):     
    return pd.concat([df1, df2]).drop_duplicates()

df3 = union(young_sailors, salty_sailors)
```

The output for `union(young_sailors, salty_sailors)` should be:

|   | sid   | sname    | rating | age |
|---|-------|----------|--------|-----|
| 0 | 2701  | Grace    | 8      | 25  |
| 1 | 18869 | Morgan   | 6      | 26  | 
| 2 | 63940 | Anne     | 4      | 21  | 
| 3 | 21869 | Jack     | 9      | 27  | 
| 4 | 17436 | Javier   | 3      | 22  | 
| 2 | 45433 | Balon    | 7      | 39  | 
| 3 | 22689 | Euron    | 10     | 35  | 
| 4 | 46535 | Victarion | 2     | 37  | 


```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Exercise 4: Intersection 

Write a function to perform the binary operation of Intersection that takes two DataFrames as input.  

You can assume the two dataframes are intersection compatible (that is, they have the same number of columns, and their columns have the same types).  

Return the result of the intersection operation.  


```python
def intersection(df1, df2):     
    return pd.merge(df1, df2, how="inner")

df4 = intersection(young_sailors, salty_sailors)

```

The output for `intersection(young_sailors, salty_sailors)` should be:

|   | sid   | sname    | rating | age |
|---|-------|----------|--------|-----|
| 0 | 2701  | Grace    | 8      | 25  |
| 1 | 17436 | Javier   | 3      | 22  | 


```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Exercise 5: Set Difference 

Write a function to perform the set difference operation that takes two DataFrames as input.  

This one is a bit harder.  

You may want to convert rows of the dataframe to a tuple, if you are having trouble. 

Alternatively, you could consider using the `isin` method.

```python
def difference(df1, df2): 
    # Assume inputs df1 and df2 are same schema 
    # Make sure your returned DataFrame has the same data types as the 
    #  input DataFrames 
    # Make sure to not change the original data frames passed as 
    #  inputs 
    
    # https://stackoverflow.com/questions/48647534/find-difference-between-two-data-frames
    return df1[~df1.apply(tuple,1).isin(df2.apply(tuple,1))]

df5 = difference(young_sailors, salty_sailors)
```

The output for `difference(young_sailors, salty_sailors)` should be:

|   | sid   | sname    | rating | age |
|---|-------|----------|--------|-----|
| 1 | 18869 | Morgan   | 6      | 26  | 
| 2 | 63940 | Anne     | 4      | 21  | 
| 3 | 21869 | Jack     | 9      | 27  | 


```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

## Exercise 6: Cross-product 

Implement a function for the cross-product between two input DataFrames

This one is also harder, so there is some code to start.  Think about how the new `tmp-key` column could be used ...

```python
def cross_product(in1, in2):
    # add a column "tmp-key" of zeros to df1 and df2 
    in1 = pd.concat([in1, pd.Series(0, index=in1.index, name="tmp-key")], axis=1)
    in2 = pd.concat([in2, pd.Series(0, index=in2.index, name="tmp-key")], axis=1)
    # use Pandas merge functionality along with drop 
    df6 = pd.merge(in1, in2, how='cross').drop(columns=["tmp-key_x", "tmp-key_y"])
    # to compute outer product and remove extra column
    return df6

df6 = cross_product(young_sailors, salty_sailors)
```

The output of `cross_product(young_sailors, salty_sailors)` should be: 

|   | sid_x | sname_x  | rating_x | age_x | sid_y  | sname_y | rating_y | age_y |
|---|-------|----------|--------|-----|-------|--------|-----|----|
| 0 | 2701  | Grace    | 8      | 25  | 2701  | Grace  | 8   | 25 |
| 1 | 2701  | Grace    | 8      | 25  | 17436 | Javier | 3   | 22 |
| 2 | 2701  | Grace    | 8      | 25  | 45433 | Balon  | 7   | 39 | 
| 3 | 2701  | Grace    | 8      | 25  | 22689 | Euron  | 10  | 35 |
| 4 | 2701  | Grace    | 8      | 25  | 46535 | Victarion | 2   | 37 |
| 5 | 18869 | Morgan   | 6      | 26  | 2701  | Grace  | 8   | 25 |
|   | ...   | ...      | ...    | ... | ...   | ...    | ... | ... |

There should be 25 entries.

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

## Exercise 7: Theta-Join 

Complete a function for theta-join.  Consider using two other relational operators you have created above.

```python
def theta_join(df1, df2, condition):
    temp = cross_product(df1, df2)
    
    return select(temp, condition)

df7 = theta_join(young_sailors, salty_sailors, lambda x: x["age_x"] > x["age_y"])
```

The output of `theta_join(young_sailors, salty_sailors, lambda x: x["age_x"] > x["age_y"])` should be: 

|   | sid_x | sname_x  | rating_x | age_x | sid_y  | sname_y | rating_y | age_y |
|---|-------|----------|--------|-----|-------|--------|-----|----|
| 0 | 2701  | Grace    | 8      | 25  | 17436 | Javier | 3   | 22 |
| 5 | 18869 | Morgan   | 6      | 26  | 2701  | Grace  | 8   | 25 |
| 6 | 18869 | Morgan   | 6      | 26  | 17436 | Javier | 3   | 22 | 
| 15| 21869 | Jack     | 9      | 27  | 2701  | Grace  | 8   | 25 |
| 16| 21869 | Jack     | 9      | 27  | 17436 | Javier | 3   | 22 |


```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

## Exercise 8: Natural Join 

Implement a function to return the natural join.  Try to implement this using other relational operators.

```python
def natural_join(df1, df2, attr):
    df8 = pd.merge(df1, df2, how="cross")
    df8 = df8[df8[ attr + '_x'] == df8[attr + '_y']]
    # df8 = pd.merge(df1, df2, on=attr)
    return df8

all_sailors = union(young_sailors, salty_sailors)
df8 = natural_join(all_sailors, reservations, "sid")
```

The output of the natural join is: 

|   | sid_x |	sname |	rating | age  | sid_y | bid | day |
|---|-------|---------|--------|------|-------|-----|-----|
| 12 |	18869 |	Morgan | 6 | 26 | 18869 | 50041 | 3/2 |
| 19 |	18869 |	Morgan | 6 | 26 | 18869 | 41116 | 3/4 |
| 26 |	63940 |	Anne | 4 | 21 |	63940 |	54505 |	3/3 |
| 30 |	21869 |	Jack | 9 | 27 |	21869 |	41116 |	3/1 |
| 34 |	21869 |	Jack | 9 | 27 |	21869 |	58324 |	3/2 |
| 38 |	21869 |	Jack | 9 | 27 |	21869 |	50041 |	3/3 |
| 45 |	17436 |	Javier | 3 | 22 |	17436 |	50041 |	3/3 |
| 51 |	45433 |	Balon |	7 |	39 |	45433 |	35168 |	3/1 |
| 57 |	45433 |	Balon |	7 |	39 |	45433 |	41116 |	3/3 |
| 63 |	22689 |	Euron |	10 | 35 |	22689 |	41116 |	3/2 |

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

## Congratulations! You have finished Lab9!

Please submit the zip file to Gradescope. 

Make sure you have run all cells in your notebook **in order** before running the cell below, so that all information gets saved to the log file correctly. The cell below will generate a zip file for you to submit. **Please save before exporting!**


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
/home/campus21/yaljnadi/.conda/envs/un5550fa24/lib/python3.10/site-packages/otter/check/notebook.py:470: UserWarning: Could not locate a PDF to include
  warnings.warn("Could not locate a PDF to include")
File ~/.conda/envs/un5550fa24/lib/python3.10/site-packages/otter/export/__init__.py:36, in export_notebook(nb_path, dest, exporter_type, **kwargs)
     33     pdf_name = os.path.splitext(nb_path)[0] + ".pdf"
     35 Exporter = get_exporter(exporter_type=exporter_type)
---> 36 Exporter.convert_notebook(nb_path, pdf_name, **kwargs)
     38 return pdf_name

File ~/.conda/envs/un5550fa24/lib/python3.10/site-packages/otter/export/exporters/via_latex.py:77, in PDFViaLatexExporter.convert_notebook(cls, nb_path, dest, xecjk, **kwargs)
     73     if xecjk:
     74         message += "\n\nIf the error above is related to xeCJK or fandol in LaTeX " \
     75             "and you don't require this functionality, try running again without " \
     76             "xecjk set to True or the --xecjk
```

```python

```

```python

```

Tests for q1 (points: [0.5, 0.5, 0.5, 0.5, 1])
```
>>> df1a.shape[0] == 5 and df1a.shape[1] == 2
True
```
```
>>> all(df1a.columns == ['sname', 'age'])
True
```
```
>>> df1a.iloc[1, 1] == 22 and df1a.iloc[3, 0] == 'Euron'
True
```
```
>>> df1b.shape[0] == 5 and df1b.shape[1] == 1
True
```
```
>>> all(df1b.columns == 'color') and df1b.iloc[2, 0] == 'Silver'
True
```


Tests for q2 (points: [2, 2])
```
>>> dfq2.shape[0] == 2 and dfq2.shape[1] == 4
True
```
```
>>> dfq2.iloc[1, 3] == 27 and dfq2.iloc[0, 2] == 8
np.True_
```


Tests for q3 (points: 8)
```
>>> df3.shape[0] == 8 and df3.shape[1] == 4
True
```


Tests for q4 (points: 8)
```
>>> intersection(young_sailors, salty_sailors).shape == (2, 4)
True
```


Tests for q5 (points: [1, 1, 1, 1])
```
>>> df5.shape == (3, 4)
True
```
```
>>> df5.iloc[2, :].to_list() == [21869, 'Jack', 9, 27]
True
```
```
>>> df5.dtypes.iloc[0] == 'int64'
True
```
```
>>> df5.dtypes.iloc[3] == 'int64'
True
```


Tests for q6 (points: [2, 2])
```
>>> df6.shape == (25, 8)
True
```
```
>>> df6.iloc[8, 3:6].to_list() == [26, 22689, 'Euron']
True
```


Tests for q7 (points: 8)
```
>>> df7.shape == (5, 8)
True
```
```
>>> df7.iloc[3, 2:5].to_list() == [9, 27, 2701]
True
```


Tests for q8 (points: 8)
```
>>> df8.shape == (10, 7)
True
```
```
>>> df8.iloc[5, 1:4].to_list() == ['Jack', 9, 27]
True
```
